## Desarrolle una función que retorne el resulado de elevar una cantidad a una potencia

<!-- grade_id: cell-6a9b79d8c6bc16bb, points: 0, grade: false, solution: true -->
```python
def elevarCuadrado(valor,exponente):
    ### BEGIN
    valorexp = valor ** exponente
    return int(valorexp) 
    ### END

#con print comprueba los casos de salida
print(elevarCuadrado(2, 3))
```

Output:
```
8
```

### Validar resultados

<!-- grade_id: cell-0d1fe3cb1c9aa2e5, points: 5, grade: true, solution: false -->
```python
assert elevarCuadrado(2, 3) == 8
assert elevarCuadrado(3, 20) == 3486784401
### BEGIN HIDDEN TESTS
assert elevarCuadrado(10, 10) == 10000000000
assert elevarCuadrado(6, 7) == 279936
assert elevarCuadrado(-6, 7) == -279936
### END HIDDEN TESTS
```

```python

```