## Desarrolle una función que retorne la media de una lista de valores

La **media** es el promedio o la suma de todos los valores dividido entre la cantidad de valores.

<!-- grade_id: cell-bc65578e77764bf5, points: 0, grade: false, solution: true -->
```python
def mediaNumeros(lista):
    ### BEGIN
    media = sum(lista)/len(lista)
    return media 
    ### END

#con print comprueba los casos de salida
print(mediaNumeros([3, 2, 6, 1, 7, -4, -12, 0, 0, 10, 1])) # = peso es normal
```

Output:
```
1.2727272727272727
```

### Validar resultados

<!-- grade_id: cell-4dc7554895f600b1, points: 5, grade: true, solution: false -->
```python
assert mediaNumeros([1, 2, 3, 4, 5, 6]) == 3.5
assert mediaNumeros([0, 1, 2, 0, 1, 2, 3, 4]) == 1.625
assert mediaNumeros([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14]) == 7.0
### BEGIN HIDDEN TESTS
assert mediaNumeros([3, 2, 6, 1, 7, -4]) == 2.5
assert mediaNumeros([3, 2, 6, 1, 7, -4, -12, 0, 0, 10, 1]) == 1.2727272727272727
### END HIDDEN TESTS
```

```python

```